---
title       : Movie Tickets
description : Insert the chapter description here

--- type:NormalExercise lang:python xp:100 skills:2 key:ccbaff9352
## Movie Tickets


*** =instructions
Angela wants to order a set of pizzas for a party. Each pizza has an original cost + cost of its toppings.
However, the receipt doesn't display the total price of all the pizzas + toppings, so you will have to help Joe
determine the price he has to pay.



*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```
